---
title       : File for the Workshop
description : Download slides for this workshop
attachments :
  slides_link : https://s3.amazonaws.com/assets.datacamp.com/course/teach/slides_example.pdf

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:b44077c753
## Files Needed for this workshop

* Slides: [https://www.slideshare.net/obscrivn/rbootcamp-day-5](https://www.slideshare.net/obscrivn/rbootcamp-day-5)


When you are done downloading files and you have studied the lecture, please complite the survey about this Rbootcamp.

* Survey: [IU Qualtrics-Survey](https://iu.co1.qualtrics.com/jfe/form/SV_6hzf2sdZkmWOuq1)

When you are ready, select DONE and proceed to your next assignment for today: the 1st chapter of the course **Text Mining Bag of Words**



*** =instructions
- DONE
- Not done yet

*** =hint


*** =pre_exercise_code
```{r}

```

*** =sct
```{r}
# SCT written with testwhat: https://github.com/datacamp/testwhat/wiki

msg1 <- "Great Job. Hope to see you again!"
msg2 <- "Click Done"
test_mc(correct = 1, feedback_msgs = c(msg1,msg2))
```

